# 2019-02-21 Exam

### General Instructions:

Welcome to the **Python Programming (for Data Science)** exam session! Please, read **carefully** the instructions below before start writing code. 

This session will last **75 minutes** and is divided into **two parts**: one about "general" Python programming and the other about Python programming for Data Science. Each part is made of a set of exercises, which globally accounts for **16** + **16** = **32 points**.
You will earn all of the points associated to an exercise **if and only if** the answer you provide passes successfully **all** the tests (both those that are visible and those that are hidden to you).<br />

To actually write down your implementation, make sure to fill in any place that says <code style="color:green">**_# YOUR CODE HERE_**</code>. Note also that you should **either comment or delete** any <code style="color:green">**raise NotImplementedError()**</code> exception.<br />

For this exam session **you will not be allowed** to use any lecture material yet you will be able to access the following APIs:

-  [Python](https://docs.python.org/3.6/library/index.html)
-  [Numpy](https://docs.scipy.org/doc/numpy-1.13.0/reference/)
-  [Scipy](https://docs.scipy.org/doc/scipy-1.0.0/reference/)
-  [Pandas](https://pandas.pydata.org/pandas-docs/version/0.22/api.html)
-  [Matplotlib](https://matplotlib.org/2.1.1/api/index.html)
-  [Seaborn](http://seaborn.pydata.org/api.html)

Once you are done, save this notebook and rename it as follows:

<code>**YOURUSERNAME_2019-02-21.ipynb**</code>

where <code>**YOURUSERNAME**</code> is your actual username. To be consistent, we are expecting your username to be composed by your first name's initial, followed by your full lastname. As an example, in my case this notebook must be saved as <code>**gtolomei_2019-02-21.ipynb**</code> (Remember to insert an underscore <code>**'_'**</code> between your username and the date).<br />

Finally, go back to the [Moodle](https://elearning.studenti.math.unipd.it/esami/mod/assign/view.php?id=475) web page of the "**2019-02-21 Python Programming Exam**"; there, you will be able to upload your notebook file for grading.

<center><h3>Submissions are allowed until <span style="color:red">Thursday, 21 February 2019 at 10:45 AM</span></h3></center>

Note that there is no limit on the number of submissions; however, be careful when you upload a new version of this notebook because each submission overwrites the previous one. 
The due date indicated above is **strict**; after that, the system will not accept any more submissions and the latest uploaded notebook will be the one considered for grading.

The archive you have downloaded (<code style="color:magenta">**2019-02-21-exam.tar**</code>) is orgaized as follows:

<code style="color:red">**2019-02-21-exam**</code> (root)<br />
|----<code style="color:green">**2019-02-21.ipynb**</code> (_this_ notebook)<br />
|----<code>**corpus.txt**</code> (the text corpus you will be using for answering general Python programming questions)<br />
|----<code>**dataset.csv**</code> (the dataset you will be using for answering data science related questions)<br />
|----<code>**README.txt**</code> (a description of the dataset above)

<center><h3>... Now, sit back, relax, and do your best!</h3></center>

**First Name** = Your _first name_ here

**Last Name** = Your _last name_ here

```python
import math
import string
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
# Adding the following line, allows Jupyter Notebook to visualize plots
# produced by matplotlib directly below the code cell which generated those.
%matplotlib inline
import seaborn as sns
from scipy.stats import pearsonr
from nose.tools import assert_equal
from operator import itemgetter

EPSILON = .0000001 # tiny tolerance for managing subtle differences resulting from floating point operations

TEXT_CORPUS_FILE = "corpus.txt"
DATASET_FILE = "dataset.csv"
```

# Part 1: General Coding (16 points)

For **Part 1**, you will be asked to use the list below - called <code>**corpus**</code> - which contains a list of text documents, where each document is represented by a lowercase string with no punctuation character whatsoever.<br /> 
Please, execute the cell right below to successfully load those documents into <code>**corpus**</code>, see a few sample documents, and then answer the following questions.

```python
# used to replace any punctuation symbol with an empty character ('')
translator = str.maketrans('', '', string.punctuation)
# load each individual document as a lowercase string into the list of strings `corpus`
corpus = [doc.strip().lower().translate(translator) for doc in open(TEXT_CORPUS_FILE)]
# print out the first 5 documents loaded
print("The following are the first 5 documents loaded out of a total of {} documents:\n".format(len(corpus)))
print("\n".join(corpus[:5]))
```

## Exercise 1.1 (1 point)

Implement the function <code>**avg_doc_length**</code>, which returns the **average length** calculated over the documents in the <code>**corpus**</code>.<br />
We define the _length_ of a document the number of the tokens which the document string is made of; a _token_ is any substring which is separated from the other by a **whitespace character**, i.e., <code>**" "**</code>.

(**EXAMPLE:** If the document you are working with is the string <code>**"I think therefore I am"**</code>, then the corresponding tokens will be: <code>**"I"**</code>, <code>**"think"**</code>, <code>**"therefore"**</code>, <code>**"I"**</code>, and <code>**"am"**</code> thereby the length of this document will be **5**.

<!-- grade_id: exercise-1-1-code, points: 0, grade: false, solution: true -->
```python
def avg_doc_length():
    """
    Returns the average length calculated over the documents in the `corpus`.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-1-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `avg_doc_length` function
"""

# Tests
assert_equal(True, avg_doc_length() < 15.6)
assert_equal(True, avg_doc_length() >= 4.2)
```

## Exercise 1.2 (3 points)

Implement the function <code>**is_word_in_doc**</code>, which takes as input a string <code>**word**</code> and and integer <code>**doc_id**</code>, and returns <code>**True**</code> if and only if the token <code>**word**</code> appears in the document <code>**doc_id**</code>. By convention, documents are identified by their index position in the <code>**corpus**</code> list, therefore the first element of the list will correspond to <code>**doc_id = 0**</code>, the second to <code>**doc_id = 1**</code>, and so on and so forth.<br />
The function must be _case-insensitive_, meaning that <code>**is_word_in_doc("galileo", 42) = is_word_in_doc("Galileo", 42) = is_word_in_doc("GALILEO", 42)**</code>. Finally, if the input <code>**doc_id**</code> is outside of its valid range $[0, N-1]$ (where $N$ is the total number of documents in the <code>**corpus**</code> list), the function should immediately return <code>**False**</code>.

(**NOTE:** Words of documents contained in the <code>**corpus**</code> list are already lowercased, but there is no restiction on the <code>**word**</code> input to the <code>**is_word_in_doc**</code> function.)

<!-- grade_id: exercise-1-2-code, points: 0, grade: false, solution: true -->
```python
def is_word_in_doc(word, doc_id):
    """
    Return True iff the string `word` appears within the document `doc_id`, False otherwise.
    Plus, it returns False whenever `doc_id` is outside of its valid range of values [0, N-1] (N = len(corpus))
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `is_word_in_doc` function
"""

# Tests
assert_equal(True, is_word_in_doc("network", 4))
assert_equal(False, is_word_in_doc("Cable", 198))
assert_equal(False, is_word_in_doc("emails", 73))
assert_equal(True, is_word_in_doc("EMAIL", 73))
```

## Exercise 1.3 (5 points)

Implement the function <code>**doc_stats**</code>, which returns a custom data structure, i.e., dictionary, where each key is a <code>**doc_id**</code> and each value is a tuple containing the <code>**min**</code>, <code>**max**</code>, <code>**mean**</code>, and <code>**standard deviation**</code> (in this very specific order) of the number of characters of the words which _that_ <code>**doc_id**</code> is made of.<br />
For example, if the document is <code>"I have been to Chargoggagoggmanchauggagoggchaubunagungamaugg lake last summer"</code>, then:
-  <code>**min = 1**</code>
-  <code>**max = 45**</code>
-  <code>**mean = 8.75**</code>
-  <code>**std_dev = 13.77**</code>

(**NOTE:** The _Chargoggagoggmanchauggagoggchaubunagungamaugg_ lake truly exists, and is located in Webster, Massachussets, USA)

<!-- grade_id: exercise-1-3-code, points: 0, grade: false, solution: true -->
```python
def doc_stats():
    """
    Returns a dictionary where each key is a `doc_id` and each value is a tuple containing 
    the min, max, mean, and standard deviation of the number of characters of each word for that document.
    """
    doc_stats = {} # This is the variable that needs to be returned
    # YOUR CODE HERE
    raise NotImplementedError()
    return doc_stats
```

<!-- grade_id: exercise-1-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `doc_stats` function
"""

# Call off the function implemented above
stats = doc_stats()

# Tests
assert_equal(2, stats[0][0])
assert_equal(11, stats[0][1])
assert_equal(True, np.abs(7.25 - stats[0][2]) < EPSILON)
assert_equal(True, np.abs(3.2691742076555053 - stats[0][3]) < EPSILON)
```

## Exercise 1.4 (7 points)

Implement the function <code>**get_most_similar_docs**</code>, which takes as input a document identifier <code>**doc_i**</code> and an integer <code><b>n</b></code> (<code>**n > 0**</code>), and returns an **ordered list of pairs**, where each pair is as follows: $(doc_j, sim_{i,j}), (j\neq i)$, i.e., the first element is a document identifier, whilst the second is the value of **Jaccard_similarity** computed between the set of $n$-grams of <code>**doc_i**</code> and <code>**doc_j**</code>. Such a list must be sorted by similarity (in not-ascending order) and by document identifier (in not-descending order).<br />
To compute Jaccard similarity between any two documents you firstly have to extract word $n$-grams out of those documents. For example, if the string document is <code>"I really like python programming"</code> then:
-  **bi-grams** ($n$ = 2): <code>**[("I", "really"), ("really", "like"), ("like", "python"), ("python", "programming")]**</code>
-  **tri-grams** ($n$ = 3): <code>**[("I", "really", "like"), ("really", "like", "python"), ("like", "python", "programming")]**</code>
-  ...

Finally, suppose $A$ and $B$ represents the sets of $n$-grams as extracted from <code>**doc_i**</code> and another document in the corpus <code>**doc_j**</code>, respectively. Then, the Jaccard similarity between $A$ and $B$ is computed as follows:
$$
J(A,B) = \frac{|A \cap B|}{|A \cup B|}
$$

**SUGGESTIONS:** Implement the following **two** helper functions: 
-  <code>**n_grams(doc, n)**</code> which takes as input the string representing a document and an integer <code><b>n</b></code> (<code>**n > 0**</code>), and extracts the $n$-grams from it, $n =$ <code>**n**</code> (pay attention to how the "sliding window" should move across the string in order to extract the corresponding substrings...)
-  <code>**jaccard_similarity(a, b)**</code> which takes as input two sets <code><b>a</b></code> and <code><b>b</b></code> and returns the Jaccard similarity as specified above (to avoid 0 division error, the function returns 0 if **both** <code><b>a</b></code> and <code><b>b</b></code> are empty.)

<!-- grade_id: exercise-1-4-code, points: 0, grade: false, solution: true -->
```python
### SUGGESTION: Implement the function `n_grams` below which takes as input a string `doc` representing a document
### and an integer n > 0, and returns a list of tuples containing the n-grams of `doc`
def n_grams(doc, n):
    n_grams = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return n_grams


### SUGGESTION: Implement the function `jaccard_similarity` below which takes as input two sets `a` and `b`
### and computes J = |a & b|/|a U b| (to avoid 0 division error, returns 0 if both a and b are empty)
def jaccard_similarity(a, b):
    if len(a) == 0 and len(b) == 0:
        return 0
    return len(set(a) & set(b))/len(set(a) | set(b))
    

def get_most_similar_docs(doc_i, n):
    """
    Returns a list of pairs (doc_j, n_grams_jaccard(doc_i, doc_j) [doc_i != doc_j] 
    sorted by n_grams_jaccard (from the highest to the lowest) and by doc_id.
    """
    most_similar_docs = [] # This is the variable that you shall return
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return most_similar_docs 
```

<!-- grade_id: exercise-1-4-test, points: 7, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_most_similar_docs` function
"""

assert_equal(249, get_most_similar_docs(569, 2)[0][0])
assert_equal(True, np.abs(0.21428571428571427 - get_most_similar_docs(569, 2)[0][1]) < EPSILON)
assert_equal(504, get_most_similar_docs(27, 1)[1][0])
assert_equal(True, np.abs(0.2 - get_most_similar_docs(27, 1)[1][1]) < EPSILON)
```

# Part 2: Data Science (16 points)

In this part, you will be working with the dataset file <code>**dataset.csv**</code>. For a complete description of this data source, please refer to the <code>**README.txt**</code> file included in the archive.
In a nutshell, this dataset contains **721** unique Pokemons, including their number (ID), name, first and second type, and basic stats: HP, Attack, Defense, Special Attack, Special Defense, and Speed. Finally, it also shows whether the Pokemon is "legendary" or not.<br />
The cell below is responsible for correctly loading the dataset from the <code>**dataset.csv**</code> file. Once this is executed, you can start answering the questions below.

```python
# Load the dataset stored at `DATASET_FILE` using "," as field separator and '?' to detect NAs
# and the specified columns as header

# column names used as header
colnames = ['id', 'name', 'type_1', 'type_2', 'total', 'hp', 'attack', 'defense', 
            'special_attack', 'special_defense', 'speed', 'generation', 'is_legendary']

# load dataset
data = pd.read_csv(DATASET_FILE, 
                   sep=',',
                   header=0,
                   names=colnames,
                   na_values='?')

# remove any duplicates
data = data.drop_duplicates('id', keep='first', inplace=False)
data.reset_index(inplace=True, drop=True)

print("Loaded `Pokemon` dataset into a dataframe of size ({} x {})".format(data.shape[0], data.shape[1]))

data.head()
```

## Exercise 2.1 (1 point)

Implement the function <code>**get_the_quickest**</code> below. This takes as input a <code>**pandas.DataFrame**</code> object, and returns the record (i.e., the <code>**pandas.Series**</code>) of the quickest Pokemon in the collection (i.e., the one with the highest <code>**speed**</code>).

<!-- grade_id: exercise-2-1-code, points: 0, grade: false, solution: true -->
```python
def get_the_quickest(data):
    """
    Returns the record of the quickest Pokemon in the collection (i.e., the one with the highest speed)
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_the_quickest` function
"""

assert_equal("Ninjask", get_the_quickest(data)['name'].iloc[0])
assert_equal(45, get_the_quickest(data)['defense'].iloc[0])
assert_equal("Bug", get_the_quickest(data)['type_1'].iloc[0])
```

## Exercise 2.2 (3 points)

Implement the function <code>**attack_stats**</code> below. This takes as input a <code>**pandas.DataFrame**</code> object and returns a tuple containing the min, max, avg, and median value of <code>**attack**</code> feature, yet computed on a _slice_ of the input <code>**pandas.DataFrame**</code>.<br />
The sliced dataset represents the subpopulation containing **legendary** Pokemons whose speed is **strictly above the overall average**, and whose defense ranges in $[52, 73)$.

<!-- grade_id: exercise-2-2-code, points: 0, grade: false, solution: true -->
```python
def attack_stats(data):
    """
    Returns a tuple containing the min, max, avg, and median value of `attack` feature,
    yet limited to a slice of the input DataFrame (data). 
    In particular, this slice will contain instances referring to legendary Pokemons
    whose speed is strictly above the overall average, and whose defense ranges in [52,73).
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `attack_stats` function
"""

# Call off `attack_stats` function
stats = attack_stats(data)

assert_equal(90, stats[0])
assert_equal(125, stats[1])
```

## Exercise 2.3 (5 points)

Implement the function <code>**get_top_k_strongest**</code> below, which takes as input a <code>**pandas.DataFrame**</code> and an integer <code><b>k</b></code>, and returns the top-<code><b>k</b></code> Pokemon's <code>**type_1**</code> with the **highest average strength**. The overall strength is provided by the <code>**total**</code> column, as it contains a summary of all the statistics described in the other available columns.

(**SUGGESTION:** In order to answer this question, you will need to compute the average strength for each group of Pokemon's <code>**type_1**</code>...)

<!-- grade_id: exercise-2-3-code, points: 0, grade: false, solution: true -->
```python
def get_top_k_strongest(data, k):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_top_k_strongest` function
"""

assert_equal(5, len(get_top_k_strongest(data, 5)))
assert_equal('Steel', get_top_k_strongest(data, 3)[1])
```

## Exercise 2.4 (7 points)

This exercise is made of **3** main questions, which you can answer independently to each other.

### Question 1 (1 point)

Feature <code>**generation**</code> represents an ordinal (numerical) variable which can take on <b>6</b> distinct values.
Assign to the variable <code>**lowest_generation**</code> below the total number of first-generation Pokemons in the dataset.

<!-- grade_id: exercise-2-4-1-code, points: 0, grade: false, solution: true -->
```python
lowest_generation = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise-2-4-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the `lowest_generation`
"""

assert_equal(False, (lowest_generation == None))
```

### Question 2 (3 points)

Plot the regression line of <code>**special_attack**</code> ($x$-axis, independent variable) against <code>**special_defense**</code> ($y$-axis, dependent variable) using <code>**sns.regplot**</code> and assign the result of the plot to the variable <code>**reg_plot**</code>. 

In addition to that, assign to the variable <code>**pearson_r**</code> the Pearson's correlation coefficient computed between the two random varibles (<code>**special_attack**</code> and <code>**special_defense**</code>). This can be computed by calling the <code>**pearsonr**</code> scipy's built-in function, which takes as input the two random variables and returns a **pair**: the first item is the Pearson's correlation coefficient, whilst the second item is the $p$-value associated to the computed statistic.

<!-- grade_id: exercise-2-4-2-code, points: 0, grade: false, solution: true -->
```python
reg_plot = None # assign this to the outcome of sns.regplot call
pearson_r = None # assign this to value of the Pearson's correlation coefficient

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise-2-4-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of `reg_plot` and `pearson_r`
"""

assert_equal(False, (reg_plot == None))
assert_equal(False, (pearson_r == None))
```

### Question 3 (3 points)

Plot the boxplot of how the variable <code>**speed**</code> is distributed across each of the <b>6</b> generations of Pokemons using <code>**sns.boxplot**</code> function. Then, assign to the list variable <code>**outliers**</code> (initially empty, i.e., <code>**outliers = []**</code>) the values of Pokemon generations which exhibit _any_ outlier. For example, if generations <b>1</b>, <b>3</b>, and <b>5</b> show some outliers then you should make the following assignment: <code>**outliers = [1, 3, 5]**</code> (order **doesn't** matter!).

(**NOTE:** If there is no outlier in any Pokemon generation then you should leave the list <code>**outliers**</code> empty as it originally is.)

<!-- grade_id: exercise-2-4-3-code, points: 0, grade: false, solution: true -->
```python
box_plot = None # assign this to the outcome of sns.boxplot call
outliers = [] # change this according to the resulting box plot!

# Create a Figure containing 1x1 subplots
fig, ax = plt.subplots(1, 1, figsize=(8,6))
# Box plot 'speed' against 'generation'
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise-2-4-3-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of `box_plot` and `outliers`
"""

assert_equal(False, (box_plot == None))
assert_equal(False, (outliers == None))
```