# 2019-07-15 Exam

### General Instructions:

Welcome to the **Python Programming (for Data Science)** exam session! Please, read **carefully** the instructions below before start writing code. 

This session will last **75 minutes** and is divided into **two parts**: one about "general" Python programming and the other about Python programming for Data Science. Each part is made of a set of exercises, which globally accounts for **16** + **16** = **32 points**.
You will earn all of the points associated to an exercise **if and only if** the answer you provide passes successfully **all** the tests (both those that are visible and those that are hidden to you).<br />

To actually write down your implementation, make sure to fill in any place that says <code style="color:green">**_# YOUR CODE HERE_**</code>. Note also that you should **either comment or delete** any <code style="color:green">**raise NotImplementedError()**</code> exception.<br />

For this exam session **you will not be allowed** to use any lecture material yet you will be able to access the following APIs:

-  [Python](https://docs.python.org/3.6/library/index.html)
-  [Numpy](https://docs.scipy.org/doc/numpy-1.13.0/reference/)
-  [Scipy](https://docs.scipy.org/doc/scipy-1.0.0/reference/)
-  [Pandas](https://pandas.pydata.org/pandas-docs/version/0.22/api.html)
-  [Matplotlib](https://matplotlib.org/2.1.1/api/index.html)
-  [Seaborn](http://seaborn.pydata.org/api.html)

Once you are done, save this notebook and rename it as follows:

<code>**YOURUSERNAME_2019-07-15.ipynb**</code>

where <code>**YOURUSERNAME**</code> is your actual username. To be consistent, we are expecting your username to be composed by your first name's initial, followed by your full lastname. As an example, in my case this notebook must be saved as <code>**gtolomei_2019-07-15.ipynb**</code> (Remember to insert an underscore <code>**'_'**</code> between your username and the date).<br />

Finally, go back to the [Moodle](https://elearning.studenti.math.unipd.it/esami/mod/assign/view.php?id=553) web page of the "**2019-07-15 Python Programming Exam**"; there, you will be able to upload your notebook file for grading.

<center><h3>Submissions are allowed until <span style="color:red">Monday, 15 July 2019 at 10:45 AM</span></h3></center>

Note that there is no limit on the number of submissions; however, be careful when you upload a new version of this notebook because each submission overwrites the previous one. 
The due date indicated above is **strict**; after that, the system will not accept any more submissions and the latest uploaded notebook will be the one considered for grading.

The archive you have downloaded (<code style="color:magenta">**2019-07-15-exam.tar**</code>) is orgaized as follows:

<code style="color:red">**2019-07-15-exam**</code> (root)<br />
|----<code style="color:green">**2019-07-15.ipynb**</code> (_this_ notebook)<br />
|----<code>**dataset.csv**</code> (the dataset you will be using for answering data science related questions)<br />
|----<code>**README.txt**</code> (a description of the dataset above)

<center><h3>... Now, sit back, relax, and do your best!</h3></center>

**First Name** = PLEASE, FILL THIS WITH YOUR _FIRST NAME_

**Last Name** = PLEASE, FILL THIS WITH YOUR _LAST NAME_

```python
import math
import copy
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
# Adding the following line, allows Jupyter Notebook to visualize plots
# produced by matplotlib directly below the code cell which generated those.
%matplotlib inline
import seaborn as sns
from collections import Counter
from nose.tools import assert_equal
from operator import itemgetter
from scipy.stats import pearsonr

EPSILON = .0000001 # tiny tolerance for managing subtle differences resulting from floating point operations

DATASET_FILE = "dataset.csv"
```

# Part 1: General Coding (16 points)

For **Part 1**, you will be asked to use the dictionary below - called <code>**friends**</code> - which represents a small social network graph extracted from **10** characters of the _Friends_ TV series. Each entry of this dictionary contains the name of the character as key and a list of tuples as value; each tuple, in turn, is composed of two items: the name of a character and a float in the range **(0, 1]** which measures the strength of this relationship.<br /> 
Please, execute the cell right below and answer the following questions.

```python
friends = {
    'Rachel Green': [('Ross Geller', 1.0), 
                     ('Paul Stevens', 0.46),
                     ('Chandler Bing', 0.73),
                     ('Monica Geller', 0.92),
                     ('Phoebe Buffay', 0.94),
                     ('Joey Tribbiani', 0.97),
                     ('Gunther', 0.01)
                    ],
    'Monica Geller': [('Ross Geller', 0.85),
                      ('Phoebe Buffay', 0.89),
                      ('Rachel Green', 0.89),
                      ('Chandler Bing', 1.0),
                      ('Joey Tribbiani', 0.90),
                      ('Janice Hosenstein', 0.07),
                      ('Gunther', 0.51),
                      ('Ursula Buffay', 0.42),
                      ('Paul Stevens', 0.27)
                     ],
    'Phoebe Buffay': [('Ursula Buffay', 0.01),
                      ('Joey Tribbiani', 0.96),
                      ('Monica Geller', 0.81),
                      ('Rachel Green', 0.88),
                      ('Chandler Bing', 0.64),
                      ('Ross Geller', 0.83),
                      ('Gunther', 0.25)
                     ],
    'Joey Tribbiani': [('Phoebe Buffay', 0.98),
                       ('Chandler Bing', 1.0),
                       ('Ross Geller', 0.99),
                       ('Rachel Green', 0.97),
                       ('Monica Geller', 0.95)
                      ],
    'Chandler Bing': [('Monica Geller', 1.0),
                      ('Joey Tribbiani', 1.0),
                      ('Ross Geller', 0.97),
                      ('Phoebe Buffay', 0.81),
                      ('Rachel Green', 0.68),
                      ('Gunther', 0.37),
                      ('Janice Hosenstein', 0.16)
                     ],
    'Ross Geller': [('Monica Geller', 0.95),
                    ('Joey Tribbiani', 0.96),
                    ('Chandler Bing', 0.96),
                    ('Phoebe Buffay', 0.91),
                    ('Rachel Green', 1.0),
                    ('Paul Stevens', 0.04)],
    'Gunther': [('Rachel Green', 1.0)],
    'Ursula Buffay': [],
    'Janice Hosenstein': [('Chandler Bing', 1.0),
                          ('Monica Geller', 0.01),
                          ('Joey Tribbiani', 0.32)
                         ],
    'Paul Stevens': [('Rachel Green', 0.88),
                     ('Ross Geller', 0.29),
                     ('Monica Geller', 0.57)
                    ]
}
```

## Exercise 1.1 (1 point)

Implement the function <code><b>are_friends_with_each_other</b></code>, which takes as input two strings (i.e., two characters) <code><b>u</b></code> and <code><b>v</b></code> and returns <code><b>True</b></code> iff <code><b>u</b></code> is friend of <code><b>v</b></code> and <code><b>v</b></code> is friend of <code><b>u</b></code>.

(**NOTE:** In general, friendship relationship may be not symmetric; we suggest to implement an auxiliary function <code><b>is_friend(x, y)</b></code>, which returns <code><b>True</b></code> iff <code><b>x</b></code> appears in the list of <code><b>y</b></code>'s friends, and use it twice...)

<!-- grade_id: exercise-1-1-code, points: 0, grade: false, solution: true -->
```python
def is_friend(x, y):
    """
    Return True iff x appears in the list of y's friends.
    """
    ### BEGIN SOLUTION
    return x in set([f[0] for f in friends[y]])
    ### END SOLUTION


def are_friends_with_each_other(u, v):
    """
    Return True iff u and v are friends with each other.
    """
    ### BEGIN SOLUTION
    return is_friend(u, v) and is_friend(v, u)
    ### END SOLUTION
```

<!-- grade_id: exercise-1-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `are_friends_with_each_other` function
"""

# Tests
assert_equal(True, are_friends_with_each_other("Rachel Green", "Ross Geller"))
assert_equal(False, are_friends_with_each_other("Phoebe Buffay", "Gunther"))
assert_equal(True, are_friends_with_each_other("Monica Geller", "Joey Tribbiani"))
### BEGIN HIDDEN TESTS
assert_equal(False, are_friends_with_each_other("Ursula Buffay", "Chandler Bing"))
assert_equal(True, are_friends_with_each_other("Gunther", "Rachel Green"))
assert_equal(False, are_friends_with_each_other("Janice Hosenstein", "Paul Stevens"))
### END HIDDEN TESTS
```

## Exercise 1.2 (3 points)

Implement the function <code><b>add_friend</b></code>, which takes as input two characters <code><b>u</b></code> and <code><b>v</b></code> and updates the social network <code><b>friends</b></code> by adding a new link between <code><b>u</b></code> and <code><b>v</b></code>, i.e., adding <code><b>v</b></code> to the list of <code><b>u</b></code>'s friends using the default <code><b>friend_score</b></code> specified as the function's keyword argument **if and only if** the following condition holds:
-  <code><b>u</b></code> exists **and** <code><b>v</b></code> is not already a friend of <code><b>u</b></code>.<br />

Otherwise, the function should return <code><b>-1</b></code>.

(**SUGGESTION:** You may want to use the auxiliary function <code><b>is_friend(x, y)</b></code> above, if you have implemented it.)

<!-- grade_id: exercise-1-2-code, points: 0, grade: false, solution: true -->
```python
def add_friend(u, v, friend_score=0.5):
    """
    Add the tuple `(v, friend_score)` to the list of u's friends iff u exists and v is not already friend of u, 
    otherwise return -1.
    """
    ### BEGIN SOLUTION
    if u not in friends or is_friend(v, u):
        return -1
    friends[u].append((v, friend_score))
    ### END SOLUTION
```

<!-- grade_id: exercise-1-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `add_friend` function
"""

# save the original `friends` social network graph
old_friends = copy.deepcopy(friends)

# Tests
add_friend("Ursula Buffay", "Ugly Naked Guy")
assert_equal(friends["Ursula Buffay"], [("Ugly Naked Guy", 0.5)])
assert_equal(-1, add_friend("Ted Mosby", "Ross Geller"))
assert_equal(-1, add_friend("Ursula Buffay", "Ugly Naked Guy"))
### BEGIN HIDDEN TESTS
add_friend("Gunther", "Monica Geller")
assert_equal(friends["Gunther"], [('Rachel Green', 1.0), ('Monica Geller', 0.5)])
assert_equal(-1, add_friend("Barney Stinson", "Marshall Eriksen"))
assert_equal(-1, add_friend("Ross Geller", "Joey Tribbiani"))
### END HIDDEN TESTS

# revert to the original `friends` social network graph
friends = old_friends
```

## Exercise 1.3 (5 points)

Implement the function <code>**node_degree_distribution**</code>, which takes as input an array of integers <code><b>degrees</b></code> (initialized by default to $[0, 1, \ldots, 9]$), and returns the node (**out-**)degree distribution computed from the graph represented by <code><b>friends</b></code>. Such a distribution is an array $d$ of the same length of <code><b>degrees</b></code>, and each component $k$ is defined as follows:
$$
d[k] = \frac{n_k}{N}
$$
where $n_k = $ number of nodes with out-degree = $k$, and $N$ is the total number of nodes of the graph.<br />

(**NOTE:** The out-degree of a node is the number of outbound links from that node. We suggest to implement the auxiliary function <code><b>node_degree_counts</b></code> below, which is responsible for populating a dictionary containing the number of nodes with a specific out-degree.)

<!-- grade_id: exercise-1-3-code, points: 0, grade: false, solution: true -->
```python
def node_degree_counts():
    """
    Return a dictionary whose keys are out-degrees and whose values are the number of nodes 
    having such an out-degree value.
    """
    node_degree_counts = {} # This is the dictionary to populate and return
    ### BEGIN SOLUTION
    for u in friends:
        node_degree_counts[len(friends[u])] = node_degree_counts.setdefault(len(friends[u]), 0) + 1
    ### END SOLUTION
    return node_degree_counts  
    
def node_degree_distribution(degrees=np.arange(10)):
    """
    Return the distribution of node's out-degree.
    """
    ### BEGIN SOLUTION
    N = len(friends.keys())
    node_degrees = node_degree_counts()

    return np.array([node_degrees[k] / N if k in node_degrees else 0 for k in degrees])
    ### END SOLUTION
```

<!-- grade_id: cell-1-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `node_degree_distribution` function
"""

# Tests
assert_equal(True, np.abs(node_degree_distribution()[0] - 0.1) < EPSILON)
assert_equal(10, len(node_degree_distribution()))
assert_equal(0, node_degree_distribution()[2]) 
### BEGIN HIDDEN TESTS
assert_equal(7, np.argmax(node_degree_distribution()))
assert_equal(True, np.abs(node_degree_distribution()[3] - 0.2) < EPSILON)
assert_equal(True, np.abs(np.sum(node_degree_distribution()) - 1.0) < EPSILON)
assert_equal(0, node_degree_distribution()[-2])
### END HIDDEN TESTS
```

## Exercise 1.4 (7 points)

In this exercise, you are asked to implement the function <code><b>filter_friends</b></code>, which takes as input the name of a character <code><b>f</b></code> and a float value <code><b>min_prob_score</b></code> representing a _threshold_. The function assumes you will be working with a variant of the original <code><b>friends</b></code> dictionary given at the beginning; more specifically, you will be asked to firstly transform <code><b>friends</b></code> into another dictionary by implementing the function <code><b>markovian_friends</b></code> below. Such a function, transforms the original dictionary in such a way that original friendship scores are considered as probability scores (i.e., a probability distribution of the original friendship scores).<br />
The <code><b>filter_friends</b></code> function will return a sub-list of <code><b>f</b></code>'s friends containing only those characters whose probability score is **greater than or equal to** the value specified by <code><b>min_prob_score</b></code>. Finally, the sub-list must be sorted by friendship score in **non-ascending order** (i.e., from the largest to the smallest) and, within that, by lexicographical order in **non-descending order**.

<!-- grade_id: exercise-1-4-code, points: 0, grade: false, solution: true -->
```python
def markovian_friends():
    """
    Return a new dictionary representing the `friends` graph, where friendship scores are considered as probabilities.
    In other words, `markovian_friends[u] = [(v, P_uv), (w, P_uw), ..., (z, P_uz)]`, where:
    P_ui = the probability of reaching i from u (i.e., the normalized friendship score between u and i).
    """
    markovian_friends = {} # This is the data structure you have to populate and return for solving this exercise
    ### BEGIN SOLUTION
    for f in friends:
        tot_friendship_score_f = np.sum([u[1] for u in friends[f]])
        markovian_friends[f] = [(u[0], u[1]/tot_friendship_score_f) for u in friends[f]]
    ### END SOLUTION
    return markovian_friends

def filter_friends(f, min_prob_score):
    """
    Return a sub-list of f's friends whose normalized friendship score is greater than or equal to the value
    indicated by `min_prob_score`. Moreover, such a sub-list must be sorted by friendship score in non-ascending order 
    (i.e., from the largest to the smallest) and, within that, by lexicographical order in non-descending order
    (NOTE: This function must firstly turn the original `friends` graph into markovian friends using the function above).
    """
    ### BEGIN SOLUTION
    mfs = markovian_friends()
    
    return sorted(sorted([mf for mf in mfs[f] if mf[1] >= min_prob_score], 
                         key=lambda x:x[0]), 
                  key=lambda x:x[1], 
                  reverse=True)
    ### END SOLUTION
```

<!-- grade_id: exercise-1-4-test, points: 7, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `filter_friends` function
"""

assert_equal("Chandler Bing", filter_friends("Joey Tribbiani", 0.2)[0][0])
assert_equal(3, len(filter_friends("Joey Tribbiani", 0.2)))
assert_equal(True, np.abs(filter_friends("Monica Geller", 0.1)[2][1] - 0.153448275862069) < EPSILON)
assert_equal("Monica Geller", filter_friends("Paul Stevens", 0.3)[1][0])
### BEGIN HIDDEN TESTS
assert_equal(2, len(filter_friends("Paul Stevens", 0.3)))
assert_equal(True, np.abs(filter_friends("Paul Stevens", 0.5)[0][1] - 0.5057471264367817) < EPSILON)
assert_equal([], filter_friends("Ursula Buffay", 0.0))
assert_equal([('Joey Tribbiani', 0.2191780821917808), 
              ('Rachel Green', 0.20091324200913244), 
              ('Ross Geller', 0.18949771689497716), 
              ('Monica Geller', 0.18493150684931509), 
              ('Chandler Bing', 0.14611872146118723), 
              ('Gunther', 0.05707762557077626), 
              ('Ursula Buffay', 0.0022831050228310505)], filter_friends("Phoebe Buffay", 0.0))
assert_equal([], filter_friends("Chandler Bing", 0.25))
### END HIDDEN TESTS
```

# Part 2: Data Science (16 points)

In this part, you will be working with the dataset file <code>**dataset.csv**</code>. For a complete description of this data source, please refer to the <code>**README.txt**</code> file included in the archive.
In a nutshell, this dataset contains **1781** unique (anonymised) URLs, along with a set of **18 features** and a **binary class label** (<code>**TYPE**</code>), which indicates whether the corresponding URL is malicious (<b>1</b>) or not (<b>0</b>).<br />
The cell below is responsible for correctly loading the dataset from the <code>**dataset.csv**</code> file. Once this is executed, you can start answering the questions below.

```python
# Load the dataset stored at `DATASET_FILE` using "," as field separator and '?' to detect NAs

data = pd.read_csv(DATASET_FILE, 
                   sep=',',
                   na_values='?')

print("Loaded `websites` dataset into a dataframe of size ({} x {})".format(data.shape[0], data.shape[1]))

data.head()
```

## Exercise 2.1 (1 point)

Implement the function <code>**get_median_url_length**</code> below. This takes as input a <code>**pandas.DataFrame**</code> object, and returns the median url length (i.e., <code><b>URL_LENGTH</b></code> field) in the dataset.

<!-- grade_id: exercise-2-1-code, points: 0, grade: false, solution: true -->
```python
def get_median_url_length(data):
    """
    Return the median url length (i.e., `URL_LENGTH` field) computed across the dataset.
    """
    ### BEGIN SOLUTION
    return data.URL_LENGTH.median()
    ### END SOLUTION
```

<!-- grade_id: exercise-2-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_median_url_length` function
"""

assert_equal(True, get_median_url_length(data) < 100)
assert_equal(True, get_median_url_length(data) > 2)
### BEGIN HIDDEN TESTS
assert_equal(49, get_median_url_length(data))
### END HIDDEN TESTS
```

## Exercise 2.2 (3 points)

Implement the function <code><b>get_highest_delta_packets</b></code> below. This takes as input a <code>**pandas.DataFrame**</code> object, and returns the record (i.e., the <code>**pandas.Series**</code>) corresponding to the URL with the **highest difference** between the number of packets sent (<code><b>SOURCE_APP_PACKETS</b></code>) and those received (<code><b>REMOTE_APP_PACKETS</b></code>) in the dataset.

(**NOTE:** You are expecting to find the 0-based integer index of the record corresponding to the one having the highest difference between the two fields...)

<!-- grade_id: exercise-2-2-code, points: 0, grade: false, solution: true -->
```python
def get_highest_delta_packets(data):
    """
    Return the record corresponding to the URL with the highest difference between 
    the number of packets sent (SOURCE_APP_PACKETS) and those received (REMOTE_APP_PACKETS) in the dataset.
    """
    ### BEGIN SOLUTION
    return data.iloc[(data.SOURCE_APP_PACKETS - data.REMOTE_APP_PACKETS).values.argmax()]
    ### END SOLUTION
```

<!-- grade_id: exercise-2-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_highest_delta_packets` function
"""

# Call the `get_highest_delta_packets` function
record = get_highest_delta_packets(data)

assert_equal("B0_639", record["URL"])
assert_equal("iso-8859-1", record["CHARSET"])
assert_equal(0, record["TYPE"])
### BEGIN HIDDEN TESTS
assert_equal(165, record.name)
assert_equal(88, record["SOURCE_APP_PACKETS"])
assert_equal(44, record["REMOTE_APP_PACKETS"])
assert_equal("CA", record["WHOIS_STATEPRO"])
### END HIDDEN TESTS
```

## Exercise 2.3 (5 points)

Implement the function <code><b>get_countries_with_largest_avg_benign_traffic</b></code> below, which takes as input a <code><b>pandas.DataFrame</b></code> and an integer <code><b>k</b></code>, and returns an **ordered list** of <code><b>k</b></code> elements, where each element of is a tuple containing the name of the country (i.e., the <code><b>WHOIS_COUNTRY</b></code> field) and the **average** of bytes transferred (i.e., the <code><b>APP_BYTES</b></code> field) computed for that country, yet **limited to only benign web traffic** (i.e., <code><b>TYPE=0</b></code>). The final list shall be ordered by such a computed average (non-ascending) and, within that, lexicographically sorted (non-descending).

(**SUGGESTION:** In order to answer this question, you will need to compute the average of bytes transferred for each country, which must be set to **0** whenever a country is not involved in _any_ benign traffic.)

<!-- grade_id: exercise-2-3-code, points: 0, grade: false, solution: true -->
```python
def get_countries_with_largest_avg_benign_traffic(data, k):
    ### BEGIN SOLUTION
    countries = data.groupby(by=["WHOIS_COUNTRY"])
    result = []
    for country_name, country_group in countries:
        avg_benign_traffic = country_group.loc[country_group.TYPE == 0, "APP_BYTES"].mean()
        if avg_benign_traffic is np.nan:
            avg_benign_traffic = 0
        result.append((country_name, avg_benign_traffic))

    return sorted(sorted(result, key=lambda x:x[0], reverse=False), key=lambda x:x[1], reverse=True)[:k]
    ### END SOLUTION
```

<!-- grade_id: exercise-2-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_countries_with_largest_avg_benign_traffic` function
"""

# Call `get_countries_with_largest_avg_benign_traffic` function
largest_avg_benign_traffic = get_countries_with_largest_avg_benign_traffic(data, 15)

assert_equal('CA', largest_avg_benign_traffic[0][0])
assert_equal('LV', largest_avg_benign_traffic[2][0])
assert_equal(True, np.abs(largest_avg_benign_traffic[0][1] - 30184.585365853658) < EPSILON)
### BEGIN HIDDEN TESTS
assert_equal(15, len(largest_avg_benign_traffic))
assert_equal('AU', largest_avg_benign_traffic[-1][0])
assert_equal('TH', largest_avg_benign_traffic[-2][0])
assert_equal(True, np.abs(largest_avg_benign_traffic[1][1] - 10490.0) < EPSILON)
assert_equal(True, np.abs(largest_avg_benign_traffic[8][1] - 1802.1679316888046) < EPSILON)
### END HIDDEN TESTS
```

## Exercise 2.4 (7 points)

This exercise is made of **3** main questions, which you can answer independently to each other.

### Question 1 (1 point)

Feature <code><b>DIST_REMOTE_TCP_PORT</b></code> represents a categorical variable which can take on <b>66</b> distinct values.
Implement the function <code><b>top_k_used_ports</b></code> below, which takes as input an integer <code><b>k</b></code> and returns the list of the <code><b>k</b></code>-most frequently used port numbers.

(**NOTE:** To access the 0-based indices of a <code><b>pandas.Series</b></code> object you can use the <code><b>.index</b></code> property associated to that object).

<!-- grade_id: exercise-2-4-1-code, points: 0, grade: false, solution: true -->
```python
def top_k_used_ports(k):
    ### BEGIN SOLUTION
    return data.DIST_REMOTE_TCP_PORT.value_counts().index[:k]
    ### END SOLUTION
```

<!-- grade_id: exercise-2-4-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the `top_k_used_ports`
"""

assert_equal(0, top_k_used_ports(10)[0])
assert_equal(3, top_k_used_ports(10)[1])
assert_equal(15, top_k_used_ports(20)[-1])
### BEGIN HIDDEN TESTS
assert_equal(1, top_k_used_ports(15)[2])
assert_equal(4, top_k_used_ports(5)[-1])
assert_equal(708, top_k_used_ports(100)[-1])
assert_equal(66, len(top_k_used_ports(100)))
### END HIDDEN TESTS
```

### Question 2 (3 points)

Assign the correct values to the three variables below, according to the following requests.

-  <code><b>nan_record_ids</b></code> should contain the index of the records of the dataframe whose values of <code><b>CONTENT_LENGTH</b></code> is <code><b>NA</b></code>;

-  <code><b>nan_ratio</b></code> should contain the proportion of records (out of all records) of the dataframe whose values of <code><b>CONTENT_LENGTH</b></code> is <code><b>NA</b></code>;

-  <code><b>data_cpy</b></code> is a copy of the original dataframe, where where <code><b>NA</b></code> values of <code><b>CONTENT_LENGTH</b></code> will be replaced with the mean calculated over non-<code><b>NA</b></code> values.

<!-- grade_id: exercise-2-4-2-code, points: 0, grade: false, solution: true -->
```python
nan_record_ids = None # assign this to the index of records whose values of `CONTENT_LENGTH` is NA
nan_ratio = None # assign this to the ratio of NA values of `CONTENT_LENGTH`
data_cpy = data.copy() # copy of the original dataframe, where NA values of `CONTENT_LENGTH` will be replaced with the mean

### BEGIN SOLUTION
nan_record_ids = data[data.CONTENT_LENGTH.isna()].index.values
nan_ratio = data.CONTENT_LENGTH.isna().sum() / data.shape[0]
data_cpy.CONTENT_LENGTH = data.CONTENT_LENGTH.fillna(data.CONTENT_LENGTH.mean())
### END SOLUTION
```

<!-- grade_id: exercise-2-4-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of `nan_ratio`, `nan_record_ids`, and `data_cpy`
"""

assert_equal(False, (nan_ratio is None))
assert_equal(False, (data_cpy is None))
assert_equal(812, len(nan_record_ids))
### BEGIN HIDDEN TESTS
assert_equal(True, np.abs(nan_ratio - 0.45592363840539024) < EPSILON)
assert_equal(5, nan_record_ids[0])
assert_equal(8, nan_record_ids[1])
assert_equal(1779, nan_record_ids[-1])
assert_equal(True, np.all(np.abs(data_cpy.loc[nan_record_ids, "CONTENT_LENGTH"] - 11726.927760577915) < EPSILON))
### END HIDDEN TESTS
```

### Question 3 (3 points)

Plot the regression line of <code><b>NUMBER_SPECIAL_CHARACTERS</b></code> ($x$-axis, independent variable) against <code><b>URL_LENGTH</b></code> ($y$-axis, dependent variable) using <code><b>sns.regplot</b></code> and assign the result of the plot to the variable <code><b>reg_plot</b></code>. 

In addition to that, assign to the variable <code><b>pearson_r</b></code> the Pearson's correlation coefficient computed between the two random varibles (<code><b>NUMBER_SPECIAL_CHARACTERS</b></code> and <code><b>URL_LENGTH</b></code>). This can be computed by calling the <code><b>pearsonr</b></code> scipy's built-in function, which takes as input the two random variables and returns a **pair**: the first item is the _actual_ Pearson's correlation coefficient, whilst the second item is the $p$-value associated to the computed statistic.

<!-- grade_id: exercise-2-4-3-code, points: 0, grade: false, solution: true -->
```python
reg_plot = None # assign this to the outcome of sns.regplot call
pearson_r = None # assign this to value of the Pearson's correlation coefficient

### BEGIN SOLUTION
# Create a Figure containing 1x1 subplots
fig, ax = plt.subplots(1, 1, figsize=(8,6))
# Regression plot line 'special_attack' vs. 'special_defense'
reg_plot = sns.regplot(x=data.NUMBER_SPECIAL_CHARACTERS, y=data.URL_LENGTH, color='#0099cc', ax=ax)
pearson_r = pearsonr(data.NUMBER_SPECIAL_CHARACTERS, data.URL_LENGTH)[0]
### END SOLUTION
```

<!-- grade_id: exercise-2-4-3-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of `reg_plot` and `pearson_r`
"""

assert_equal(False, (reg_plot is None))
assert_equal(False, (pearson_r is None))
### BEGIN HIDDEN TESTS
assert_equal(True, np.abs(pearson_r - 0.9179859625469611) < EPSILON)
### END HIDDEN TESTS
```