# 2019-01-23 Exam

### General Instructions:

Welcome to the **Python Programming (for Data Science)** exam session! Please, read **carefully** the instructions below before start writing code. 

This session will last **75 minutes** and is divided into **two parts**: one about "general" Python programming and the other about Python programming for Data Science. Each part is made of a set of exercises, which globally accounts for **16** + **16** = **32 points**.
You will earn all of the points associated to an exercise **if and only if** the answer you provide passes successfully **all** the tests (both those that are visible and those that are hidden to you).<br />

To actually write down your implementation, make sure to fill in any place that says <code style="color:green">**_# YOUR CODE HERE_**</code>. Note also that you should **either comment or delete** any <code style="color:green">**raise NotImplementedError()**</code> exception.<br />

For this exam session **you will not be allowed** to use any lecture material yet you will be able to access the following APIs:

-  [Python](https://docs.python.org/3.6/library/index.html)
-  [Numpy](https://docs.scipy.org/doc/numpy-1.13.0/reference/)
-  [Scipy](https://docs.scipy.org/doc/scipy-1.0.0/reference/)
-  [Pandas](https://pandas.pydata.org/pandas-docs/version/0.22/api.html)
-  [Matplotlib](https://matplotlib.org/2.1.1/api/index.html)
-  [Seaborn](http://seaborn.pydata.org/api.html)

Once you are done, save this notebook and rename it as follows:

<code>**YOURUSERNAME_2019-01-23.ipynb**</code>

where <code>**YOURUSERNAME**</code> is your actual username. To be consistent, we are expecting your username to be composed by your first name's initial, followed by your full lastname. As an example, in my case this notebook must be saved as <code>**gtolomei_2019-01-23.ipynb**</code> (Remember to insert an underscore <code>**'_'**</code> between your username and the date).<br />

Finally, go back to [Moodle](https://elearning.studenti.math.unipd.it/esami/mod/assign/view.php?id=454) and check for the "**2019-01-23 Python Programming Exam**" item; there, you will be able to upload your notebook file for grading.

<center><h3>Submissions are allowed **until <span style="color:red">Wednesday, 23 January 2019 at 10:45 AM</span>**</h3></center>

Note that there is no limit on the number of submissions; however, be careful when you upload a new version of this notebook because each submission overwrites the previous one. 
The due date indicated above is **strict**; after that, the system will not accept any more submissions and the latest uploaded notebook will be the one considered for grading.

The archive you have downloaded (<code style="color:magenta">**2019-01-23-exam.tar**</code>) is orgaized as follows:

<code style="color:red">**2019-01-23-exam**</code> (root)<br />
|----<code style="color:green">**2019-01-23.ipynb**</code> (_this_ notebook)<br />
|----<code>**dataset.csv**</code> (the dataset you will be using for answering data science related questions)<br />
|----<code>**README.txt**</code> (a description of the dataset above)

<center><h3>... Now, sit back, relax, and do your best!</h3></center>

**First Name** = Your _first name_ here

**Last Name** = Your _last name_ here

```python
import math
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
# Adding the following line, allows Jupyter Notebook to visualize plots
# produced by matplotlib directly below the code cell which generated those.
%matplotlib inline
import seaborn as sns
from scipy.stats import norm
from nose.tools import assert_equal
from operator import itemgetter

EPSILON = .0000001 # tiny tolerance for managing subtle differences resulting from floating point operations

DATASET_FILE = "dataset.csv"
```

# Part 1: General Coding (16 points)

For **Part 1**, you will be asked to use the _nested_ dictionary below - called <code>**ratings**</code> - which contains ratings given to movies by a small set of **7 users**.<br /> 
Please, execute the cell right below and answer the following questions.

```python
ratings = {
    'Lisa Rose': {
        'Lady in the Water': 2.5,
        'Snakes on a Plane': 3.5,
        'Just My Luck': 3.0,
        'Superman Returns': 3.5,
        'You, Me and Dupree': 2.5,
        'The Night Listener': 3.0,
    },
    'Gene Seymour': {
        'Lady in the Water': 3.0,
        'Snakes on a Plane': 3.5,
        'Just My Luck': 1.5,
        'Superman Returns': 5.0,
        'The Night Listener': 3.0,
        'You, Me and Dupree': 3.5,
    },
    'Michael Phillips': {
        'Lady in the Water': 2.5,
        'Snakes on a Plane': 3.0,
        'Superman Returns': 3.5,
        'The Night Listener': 4.0,
    },
    'Claudia Puig': {
        'Snakes on a Plane': 3.5,
        'Just My Luck': 3.0,
        'The Night Listener': 4.5,
        'Superman Returns': 4.0,
        'You, Me and Dupree': 2.5,
    },
    'Mick LaSalle': {
        'Lady in the Water': 3.0,
        'Snakes on a Plane': 4.0,
        'Just My Luck': 2.0,
        'Superman Returns': 3.0,
        'The Night Listener': 3.0,
        'You, Me and Dupree': 2.0,
    },
    'Jack Matthews': {
        'Lady in the Water': 3.0,
        'Snakes on a Plane': 4.0,
        'The Night Listener': 3.0,
        'Superman Returns': 5.0,
        'You, Me and Dupree': 3.5,
    },
    'Toby Jones': {
        'Snakes on a Plane': 4.5, 
        'You, Me and Dupree': 1.0,
        'Superman Returns': 4.0},
}
```

## Exercise 1.1 (1 point)

Implement the function <code>**has_rated**</code> which takes as input two strings, i.e., <code>**user**</code> and <code>**movie**</code>, and returns <code>**True**</code> iff <code>**ratings**</code> contains a rating value for <code>**movie**</code> provided by <code>**user**</code>. If <code>**user**</code> has not provided _any_ rating at all, the function should **return** (i.e., not print!) the following string message: <code style="color:red">**'User {user} has not rated any movie'**</code>, where <code>**{user}**</code> stands for the actual value referenced by <code>**user**</code>. For example, if <code>**user = 'John Smith'**</code> this will be: <code style="color:red">**'User John Smith has not rated any movie'**</code>.

<!-- grade_id: exercise-1-1-code, points: 0, grade: false, solution: true -->
```python
def has_rated(user, movie):
    """
    Return True iff user has rated movie in the database of ratings,
    otherwise print the following message: 'User {user} has not rated any movie'.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-1-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `has_rated` function
"""

# Tests
assert_equal(True, has_rated('Jack Matthews', 'Snakes on a Plane'))
assert_equal(False, has_rated('Claudia Puig', 'Jaws'))
assert_equal('User Bill Gates has not rated any movie', has_rated('Bill Gates', 'Lady in the Water'))
```

## Exercise 1.2 (3 points)

Implement the function <code>**avg_user_rating**</code>, which takes as input a user <code>**user**</code> and returns the average rating provided by _that_ user. If the user has not rated any movie (i.e., if he/she is not in the database of ratings) the function will return <code>**-1**</code>.

<!-- grade_id: exercise-1-2-code, points: 0, grade: false, solution: true -->
```python
def avg_user_rating(user):
    """
    Return the average rating for this user.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `avg_user_rating` function
"""

# Tests
assert_equal(True, (np.abs(avg_user_rating('Lisa Rose') - 3.0)) < EPSILON)
assert_equal(True, (np.abs(avg_user_rating('Jack Matthews') - 3.7)) < EPSILON)
assert_equal(True, (np.abs(avg_user_rating('Claudia Puig') - 3.5)) < EPSILON)
```

## Exercise 1.3 (5 points)

Implement the function <code>**movie_rating_stats**</code>, which returns a custom data structure, i.e., dictionary, where each key is a movie and each value is a tuple containing the <code>**min**</code>, <code>**max**</code>, <code>**avg**</code>, and <code>**median**</code> values (in this very specific order), computed across users who have rated _that_ movie.

<!-- grade_id: exercise-1-3-code, points: 0, grade: false, solution: true -->
```python
def movie_rating_stats():
    """
    Return a dictionary where each key is a movie and each value is a tuple containing 
    the min, max, avg, and median values of that movie.
    """
    movie_rating_stats = {} # This is the variable that needs to be returned
    # YOUR CODE HERE
    raise NotImplementedError()
    return movie_rating_stats
```

<!-- grade_id: exercise-1-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `movie_rating_stats` function
"""

# Call off the function implemented above
stats = movie_rating_stats()

# Tests
assert_equal(2.5, stats['Lady in the Water'][0]) # minimum rating of 'Lady in the Water'
assert_equal(5.0, stats['Superman Returns'][1]) # maximum rating of 'Superman Returns'
assert_equal(True, np.abs(3.7142857142857144 - stats['Snakes on a Plane'][2]) 
             < EPSILON) # average rating of 'Snakes on a Plane'
assert_equal(3.0, stats['The Night Listener'][3]) # median rating of 'The Night Listener'
```

## Exercise 1.4 (7 points)

Implement the function <code>**get_most_similar_users**</code>, which takes as input a user <code>**user**</code> and returns an **ordered list** of most similar users to <code>**user**</code>. Similarity between two users $u$ and $v$ is defined by means of <i>cosine_similarity</i>, which is computed as follows:

$$
\textit{cosine_similarity}(u, v) = \frac{{\bf u} \cdot {\bf v}}{||{\bf u}||~||{\bf v}||} = \frac{\sum_{i=1}^n u_i~v_i}{\sqrt{\sum_{i=1}^n u_i^2}~\sqrt{\sum_{i=1}^n v_i^2}}
$$

where ${\bf u}=(u_1, u_2, \ldots, u_n)$ is the $n$-dimensional vector associated with the user $u$, and $||{\bf u}||$ is the **Euclidean norm** of the vector ${\bf u}$. More specifically, this measures the _cosine_ of the angle between the two vectors ${\bf u}$ and ${\bf v}$, therefore ranging in $[-1, 1]$. The highest the cosine similarity, the larger is the similarity between users.<br />
The number of dimensions $n$ of each user vector correspond to the total number of movies in the database of ratings, and each entry is evaluated as follows:

$$
u_i =  
\begin{cases}
    ratings[u][i],& \text{if user } u\text{ has rated movie }i\\
    0,              & \text{otherwise}
\end{cases}
$$

For a given input user $u$, this function will return a list $L(u)$, where each element is a pair $(v, \textit{cosine_similarity}(u, v))$, such that $v \neq u$. In addition, $L(u)$ must be ranked by cosine similarity in **not ascending order** (i.e., from the highest to the lowest), and if two or more users have the same cosine similarity those should appear in alphabetical order.

(**SUGGESTIONS:** Implement the helper function <code>**cosine_similarity(u, v)**</code> which computes the cosine similarity between two vectors, as prepared for you below. Note that only non-zero components of <code>**u**</code> and <code>**v**</code> contribute to the dot product of the two vectors...)

<!-- grade_id: exercise-1-4-code, points: 0, grade: false, solution: true -->
```python
### SUGGESTION: Implement the function `cosine_similarity` below which will return the cosine similarity between two input vectors.
### If you don't want to use this approach, please remember to delete this helper function before running the cell!
def cosine_similarity(u, v):
    # YOUR CODE HERE
    raise NotImplementedError()
    

def get_most_similar_users(user):
    """
    Return a list of pairs (v, cosine_similarity(user, v) [v != user] 
    sorted by cosine similarity (from the highest to the lowest) and by user name.
    """
    most_similar_users = [] # This is the variable that you shall return
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return most_similar_users 
```

<!-- grade_id: exercise-1-4-test, points: 7, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `get_most_similar_users` function
"""

assert_equal('Michael Phillips', get_most_similar_users('Jack Matthews')[3][0])
assert_equal(True, np.abs(0.8847891964365634 - get_most_similar_users('Gene Seymour')[3][1]) < EPSILON)
assert_equal(('Jack Matthews', 0.80555004051484014), get_most_similar_users('Toby Jones')[0])
```

# Part 2: Data Science (16 points)

In this part, you will be working with the dataset file <code>**dataset.csv**</code>. For a complete description of this data source, please refer to the <code>**README.txt**</code>.
In a nutshell, this is a sample of **1,000 instances** from a larger dataset containing census and demographic information for a set of citizens.
Each citizen (i.e., row in the file) is described by **14** features (i.e., columns) and labeled with an binary value (i.e., the 15th and last column called <code>**income_greater_than_50k**</code>), which takes on two values: **-1** (indicating the citizen's yearly income is **below or equal than** 50k dollars) or **+1**, otherwise.<br />
The cell below is responsible for correctly loading the dataset from the <code>**dataset.csv**</code> file. Once this is executed, you can start answering the questions below.

```python
# Load the dataset stored at `DATASET_FILE` using "," as field separator and '?' to detect NAs
data = pd.read_csv(DATASET_FILE, 
                   sep=",", 
                   na_values='?')
data.head()
```

## Exercise 2.1 (1 point)

Implement the function <code>**avg_capital_gain_per_age**</code> below. This takes as input a <code>**pandas.DataFrame**</code> object and an integer value <code>**age**</code>, and returns the **average** value of <code>**capital_gain**</code> of citizens having that age.

(**NOTE:** If the input age is not present, the <code>**mean()**</code> function is robust and returns <code>**np.nan**</code>)

<!-- grade_id: exercise-2-1-code, points: 0, grade: false, solution: true -->
```python
def avg_capital_gain_per_age(data, age):
    """
    Return the average value of capital gain for citizens having a specific age.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `avg_capital_gain_per_age` function
"""

assert_equal(True, np.abs(192.80769230769232 - avg_capital_gain_per_age(data, 32)) < EPSILON)
assert_equal(0.0, avg_capital_gain_per_age(data, 27))
assert_equal(2330.0, avg_capital_gain_per_age(data, 48))
```

## Exercise 2.2 (3 points)

Implement the function <code>**hours_per_week_stats**</code> below. This takes as input a <code>**pandas.DataFrame**</code> object and returns a tuple containing the min, max, avg, and median value of <code>**hours_per_week**</code>, yet computed on a _slice_ of the input <code>**pandas.DataFrame**</code>.<br />
The sliced dataset represents the subpopulation containing **female** citizens whose age is **between 36 and 49 years old** (extremes included), who were born in the **United-States** and having a value of <code>**education_num**</code> **strictly greater than** the _overall_ average.

<!-- grade_id: exercise-2-2-code, points: 0, grade: false, solution: true -->
```python
def hours_per_week_stats(data):
    """
    Return a tuple containing the min, max, avg, and median value of `hours_per_week` feature,
    yet limited to a slice of the input DataFrame (data). 
    In particular, this slice will contain instances referring to female citizens
    whose age is between 36 and 49 (extremes included) born in the United-States
    and having a value of education_num strictly above the overall population average.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-2-test, points: 3, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `hours_per_week_stats` function
"""

# Call off `hours_per_week_stats` function
stats = hours_per_week_stats(data)

assert_equal(5, stats[0])
assert_equal(60, stats[1])
```

## Exercise 2.3 (5 points)

Implement the function <code>**standardize_dataset**</code>, which takes as input a <code>**pandas.DataFrame**</code> object and modifies it by creating **new** columns corresponding to the **standardized** values of _numerical_ columns only, except for the target column, i.e., the very last column named <code>**income_greater_than_50k**</code>.<br />
To standardize a column $X$, each original value $x\in X$ should be modified with $x_{\text{std}}$ as follows:

$$
x_{\text{std}} = \frac{x - \mu_X}{\sigma_X}
$$

where $\mu_X$ and $\sigma_X$ are the **mean** and **standard deviation** computed across $X$, respectively.

To be sure standardization is applied to a numerical column only, you can use the following built-in function:

<code>**pd.api.types.infer_dtype(data[column], skipna=True)**</code>, which may return either <code>**'string'**</code> or <code>**'integer'**</code>, in our very specific case.

Finally, the naming convention for standardized columns is as follows: if <code>**column**</code> is the name of the original column, then its standardized version will be <code>**column_std**</code>.

<!-- grade_id: exercise-2-3-code, points: 0, grade: false, solution: true -->
```python
def standardize_dataset(data):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: exercise-2-3-test, points: 5, grade: true, solution: false -->
```python
"""
Test the correctness of the implementation of the `standardize_dataset` function
"""

# Call off `standardize_dataset` function on a copy of our original dataset
data_cpy = data.copy()
standardize_dataset(data_cpy)

assert_equal(1000, data_cpy.shape[0])
assert_equal(21, data_cpy.shape[1])
assert_equal(True, np.abs(0.093258105603949315 - data_cpy.iloc[42]['age_std']) < EPSILON)
assert_equal(True, np.abs(-0.21956059890921917 - data_cpy.iloc[73]['capital_loss_std']) < EPSILON)
```

## Exercise 2.4 (7 points)

This exercise is made of **3** main questions, which you can answer independently to each other.

### Question 1 (1 point)

Feature <code>**education_num**</code> represents an ordinal (numerical) variable which can take on **16** distinct values.
Assign to the variable <code>**highest_education**</code> below the total number of citizens in the dataset exhibiting the highest education degree which corresponds to the value **16**.

<!-- grade_id: exercise-2-4-1-code, points: 0, grade: false, solution: true -->
```python
highest_education = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise-2-4-1-test, points: 1, grade: true, solution: false -->
```python
"""
Test the correctness of the `highest_education`
"""

assert_equal(False, (highest_education == None))
```

### Question 2 (2 points)

Plot the distribution of the numerical feature <code>**age**</code> using <code>**distplot**</code> and assign the result of the plot to the variable <code>**dist_plot_age**</code>. In addition to that, assign to the variables <code>**mu**</code> and <code>**sigma**</code> the result of fitting the data to a Normal distribution. This can be achieved using the <code>**norm.fit**</code> scipy's built-in function on the observed data of interest.

<!-- grade_id: exercise-2-4-2-code, points: 0, grade: false, solution: true -->
```python
dist_plot_age = None
mu = None
sigma = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: exercise-2-4-2-test, points: 2, grade: true, solution: false -->
```python
"""
Test the correctness of `dist_plot_age`, `mu` and `sigma`
"""

assert_equal(False, (dist_plot_age == None))
assert_equal(False, (mu == None))
assert_equal(False, (sigma == None))
```

### Question 3 (4 points)

Implement the function <code>**n_left_hpw_outliers**</code> below, which takes as input a <code>**pandas.DataFrame**</code> object and returns the number of left outliers of the column <code>**hours_per_week**</code> (if any).<br />
Any data point **less than** (resp., greater than) of the left (resp., right) fence is considered an outlier. Both left and right fences are empirically computed as follows:

$$
F_\textrm{left} = Q_1 - 1.5 * \texttt{IQR};~~F_\textrm{right} = Q_3 + 1.5 * \texttt{IQR}
$$

where $Q_1$ and $Q_3$ represents the 1st and 3rd quartile of the distribution of interest **without considering NAs**, and $\texttt{IQR} = Q_3 - Q_1$.

(**SUGGESTIONS:** Start from drawing the box plot and visually check whether there is any outlier or not. You can either invoke the <code>**quantile**</code> function defined on a <code>**pandas.Series**</code> object **or** use the <code>**numpy.percentile**</code> function which takes as input a <code>**pandas.Series**</code> object or, more generally, any object that can easily be converted into a <code>**numpy.array**</code>).

<!-- grade_id: exercise-2-4-3-code, points: 0, grade: false, solution: true -->
```python
# SUGGESTION: visualize the box plot first!
# Create a Figure containing 1x1 subplots
fig, ax = plt.subplots(1, 1, figsize=(8,6))
# Box plot 'hours_per_week'
# YOUR CODE HERE
raise NotImplementedError()

def n_left_hpw_outliers(data):
    """
    Return the number of left outliers of the column `hours_per_week`
    """
    hpw_fence_left = None # value of the left fence
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: exercise-2-4-3-test, points: 4, grade: true, solution: false -->
```python
"""
Test the correctness of the `n_left_hpw_outliers`
"""

# Call off the function `n_left_hpw_outliers`
outliers = n_left_hpw_outliers(data)

assert_equal(False, outliers==None)
assert_equal(True, outliers > 100)
```